```python
# Initialize Otter
import otter
grader = otter.Notebook("lab0.ipynb")
```

---

<h3><center>E178/ME276DS - Statistics and Data Science for Engineers</center></h3>

<h2><center>Lab 0<br></center></h2>

<h1><center>SciPy and Pandas<br></center></h1>

---

```python
from resources.hashutils import *
from scipy import stats
import matplotlib.pyplot as plt
import numpy as np
from pathlib import Path
```

# Part 1) SciPy

[SciPy](https://docs.scipy.org/doc/scipy/tutorial/index.html#user-guide) provides implementations of many algorithms that are useful to engineers and scientists.  These include algorithms for:

+ numerical differentiation and integration,
+ interpolation,
+ signal processing,
+ linear algebra,
+ optimization,
+ probability and statistics,
+ and many others...

<b>The [stats](https://docs.scipy.org/doc/scipy/tutorial/stats.html) subpackage</b>

SciPy's support for probability and statistics resides in the **`stats`** subpackage. Here you can find implementations of a large number of distribution families, including all of the ones that we cover in this class. 

Below are links to the documentation for the individual distribution classes:

+ Bernoulli: [`stats.bernoulli`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.bernoulli.html#scipy.stats.bernoulli)
+ Binomial: [`stats.binom`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.binom.html#scipy.stats.binom)
+ Exponential: [`stats.expon`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.expon.html#scipy.stats.expon)
+ Continuous uniform: [`stats.uniform`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.uniform.html#scipy.stats.uniform)
+ Discrete uniform: [`stats.randint`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.randint.html#scipy.stats.randint)
+ Gaussian: [`stats.norm`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html#scipy.stats.norm)

<br>

<b> Constructors </b>

All of these classes have a common design. You can create an instance of probability distribution by passing the appropriate parameter values to the class constructor.  Here are the parameter values and their translation into our notation. 

| Distribution       | Our notation                | SciPy constructor          | Interpretation of parameters  |
|--------------------|-----------------------------|----------------------------|-------------------------------|
| Bernoulli          | $\mathcal{B}(s)$            | `stats.bernoulli(p)`       | `p`=$s$                       |
| Binomial           | $\mathcal{B}in(n,s)$        | `stats.binom(n,p)`         | `n`=$n$, `p`=$s$              |
| Exponential        | $\mathcal{E}(\lambda)$      | `stats.expon(scale)`       | `scale`=$1/\lambda$           |
| Continuous uniform | $\mathcal{U}^c(a,b)$        | `stats.uniform(loc,scale)` | `loc`=$a$, `scale`=$b-a$      |
| Discrete uniform   | $\mathcal{U}^d(a,b)$        | `stats.randint(low,high)`  | `low`=$a$, `high`=$b+1$       |
| Normal             | $\mathcal{N}(\mu,\sigma^2)$ | `stats.norm(loc,scale)`    | `loc`=$\mu$, `scale`=$\sigma$ |

<br>

<b> Methods </b>

The constructed object then has these methods:

| Method      | Returns                                                             |
|-------------|---------------------------------------------------------------------|
| `pdf(x)`    | Probability density function $p(x)$ (continuous distributions only) |
| `pmf(x)`    | Probability mass function $p(x)$ (discrete distributions only)      |
| `cdf(x)`    | Cumulative distribution function $F(x)$                             |
| `ppf(u)`    | Inverse cdf $F^{-1}(u)$                                             |
| `rvs(size)` | Draw `size` samples.                                                |
| `mean()`    | Mean of the distribution.                                           |
| `var()`     | Variance of the distribution.                                       |
| `std()`     | Standard deviation of the distribution.                             |


Next we will see how we can use these classes to perform simple probabilistic computations. 

## 1.1) Simulate the tossing of a fair coin

1. Use SciPy to create a model that encodes "heads" and "tails" as 1 and 0. Name the SciPy object `X`. 
2. Simulate 1000 tosses of the coin and store the result in a 1D NumPy array called `tosses`.
3. Use Matplotlib to make a [histogram](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.hist.html) of the result.

```python
...
```

```python
grader.check("p1p1")
```

## 1.2) Simulate the throw of a pair of dice

1. Use SciPy to create two models of a six-sided die. Call these `D1` and `D2`.
2. Simulate 1000 rolls of the two dice. Record the total dots in each roll. For example if in one roll you get a 3 from D1 and a 2 from D2, record a 5. Save these 1000 numbers as a 1D NumPy array called `rolls`.
3. Create a histogram of `rolls`. 
    + Use the `bins` argument of `hist` to create bins for each of the integers from 2 through 12 inclusive.  
    + Use the `rwidth` argument of `hist` to create some space between the bars (e.g. `rwidth=0.8`)
    + Your final plot should look something like this:

<center>
<img src="resources/dicehist.png" alt="Alt text" width="300" />
</center>

```python
...
```

```python
grader.check("p1p2")
```

# Part 2) Pandas

[Pandas](https://pandas.pydata.org/docs/user_guide/index.html) is an open-source library that specializes in the manipulation and analysis of **tabular data**. 

A **table** is a data type that consists of *columns* and *rows*, as shown in the figure below. 

<center>
<img src="resources/table.png" alt="Alt text" width="700" />
</center>

+ A table has $D$ **columns**. 
    + Each column has a **header**, which is a title that identifies the column
    + Columns are *homeogeneous* in the sense that all of the data contained in a column is of the same type. A column can contain numbers or strings, but not both. 

+ A table has $N$ **rows**. 
    + The first row contains the first element of each column, the second row contains the second element of each column, and so forth. 
    + Each row is identfied by its unique **index** (k1 through kN in the figure). 

In pandas, a table is called a [**DataFrame**](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.html), and each column is a [**Series**](https://pandas.pydata.org/docs/reference/api/pandas.Series.html). A DataFrame is a sequence of Series plus an index. 

Next we will go over some of the important things to know when using pandas. 

First, pandas is usually imported under the alias `pd`:

```python
import pandas as pd
```

## Part 2.1) Loading data from a file

Tabular data is often stored in [**csv**](https://en.wikipedia.org/wiki/Comma-separated_values) (comma-separate values) files. These can be loaded into Python using pandas' [**pd.read_csv**](https://pandas.pydata.org/docs/reference/api/pandas.read_csv.html) function. You will see in the linked documentation that this function takes many arguments. It is usually sufficient to just pass the file name as a string or as a `pathlib` object.

We demonstrate this next.

Load data from the file called `concrete_strength.csv` located in the `resources` folder under your current working directory. Store the result in the variable name `T`.

```python
T = ...
```

```python
T
```

```python
grader.check("p2p1")
```

## Part 2.2) Get the dimensions and types in a table

The following attributes and methods attached to the DataFrame object are useful for retrieving basic information about the table. 

+ [**`T.shape`**](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.shape.html): Similar to NumPy's shape attribute, it returns the number of rows and columns as a tuple. 
+ [**`T.dtypes`**](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.dtypes.html) Returns a Series listing the data type of each column.
+ [**`T.info()`**](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.info.html): Prints a summary with total row/column counts, column names, non-null value counts, data types, and memory usage.


Use these methods to do the following:
1. Store the number of rows in `T` into a variable named `N`.
2. Store the number of columns in `T` into a variable named `D`.
3. How many columns are integer-valued and how many are real-valued? Store these two numbers as `num_int` and `num_float`. 
4. What is the total amount of memory occupied by `T` in KB? Store this value as a float in `T_memory`.

**Note**:
+ For 3. and 4 you do not have to write code to automatically compute these values. You can just assign these manually based on what you learn from `T.dtypes` and `T.info()`.

```python
...
```

```python
grader.check("p2p2")
```

## Part 2.3) Select a column by its header

We can extract one or more columns from a DataFrae with the following syntax. 

Suppose `X` is a DataFrame with columns `'H1'`, `'H2'`, `'H3'`, etc. Then,

+ `X['H1']` returns a Series object with a single column `'H1'`.
+ `X[['H1','H2']]` returns a DataFrame with two columns `'H1'` and `'H2'`. (Notice that here we have passed in a list of strings).

Try it out!

1. Select the `Cement` column and assig it to the variable `cement`
2. Select the `Concrete compressive strength` column and assig it to the variable `strength`
3. Make a scatter plot of cement vs. strength.

```python
...
```

```python
grader.check("p2p3")
```

## Part 2.4) Save to a file

Pandas supports reading and writing tables in a number of formats, including csv, json, xml, MS Excel, and others. In this class we will use the simplest of these, which is the csv file. To save a DataFrame to a csv file, use the [**`to_csv`**](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.to_csv.html) method.

Like `read_csv`, the `to_csv` method has many arguments that can be used to configure the output. The simplest way to use it is just to pass a file name, either as a string or as a `pathlib` object. 

Try it out! 

1. Select columns `Fly Ash`, `Water`, and `Superplasticizer` from DataFrame `T` and keep them in a variable named `Tcols`.
2. Save `Tcols` to a file called "Tcols.csv" in your current working directory.

```python
...
```

```python
grader.check("p2p4")
```

## Part 2.5) Select a row by its index

Row selection in pandas uses the `loc` attribute. `loc` is a dictionary-like structure that is attached to every DataFrame. The set of keys of `loc` is the index of the DataFrame. If you want to select a particular row, you pass the row's index to `loc`. 

For example, to select the row with index 1 and assign it to the variable name `row1` use:

```python 
row1 = T.loc[1]
```
One can also select multiple rows by passing a list of keys to `loc`.

```python 
rows = T.loc[[1,2,3]]
```
`loc` also accepts a second argument for column selection. This can be a single column header or a list of column headers. 

```python 
A = T.loc[[1,2,3],'Cement']
B = T.loc[[1,2,3],['Cement','Blast Furnace Slag']]
```

Finally, one can pass a **boolean mask** to `loc` in order to select rows according to some criterion. A boolean mask is a list (or array, or any sequence) of booleans with as many entries as there are rows in the table. Passing a boolean mask to `loc` selects those rows whose corresponding entry in the mask is `True`. 

We can use boolean masks, for example, to select the concrete mixtures whose compressive strength is greater than 50 MPa. 

```python 
mask = T['Concrete compressive strength']>50   % this produces a Series of booleans
strong_mixtures = T[mask]
```

Try it out!

1. Select those rows of `T` whose `Superplasticizer` content is greater than 10. Store this as `A`.
2. Within `A`, select the rows whose `Water` content is less than 150. Store this as `B`.
3. Get the number of rows that satisfies both of these criteria and save it to the variable name `Q`.

```python
...
```

```python
grader.check("p2p5")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Please submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for p1p1 (points: None)
```
>>> def get_dist_name(rv):
...     return rv.dist.name if hasattr(rv, 'dist') else rv.name
>>> 'X' in globals() and get_dist_name(X) == 'bernoulli'
True
```
```
>>> 'tosses' in globals() and isinstance(tosses, np.ndarray) and (tosses.shape == (1000,))
True
```
```
>>> len(np.unique(tosses)) == 2 and 1 in tosses and (0 in tosses)
True
```


Tests for p1p2 (points: None)
```
>>> def get_dist_name(rv):
...     return rv.dist.name if hasattr(rv, 'dist') else rv.name
>>> 'D1' in globals() and get_dist_name(D1) == 'randint'
True
```
```
>>> def get_dist_name(rv):
...     return rv.dist.name if hasattr(rv, 'dist') else rv.name
>>> 'D2' in globals() and get_dist_name(D2) == 'randint'
True
```
```
>>> 'rolls' in globals() and isinstance(rolls, np.ndarray) and (rolls.shape == (1000,))
True
```
```
>>> len(np.unique(rolls)) == 11 and np.all(np.unique(rolls) <= 12) and np.all(np.unique(rolls) >= 2)
np.True_
```


Tests for p2p1 (points: None)
```
>>> 'T' in globals() and isinstance(T, pd.core.frame.DataFrame)
True
```
```
>>> T.shape == (1030, 9)
True
```


Tests for p2p2 (points: None)
```
>>> get_hash(N, 4) == '801fd17457152851d73cb69c3e286a7a'
True
```
```
>>> get_hash(D, 2) == 'cf5f238fca8b6e155078aa41c175743a'
True
```
```
>>> get_hash(num_int, 2) == 'e4c2e8edac362acab7123654b9e73432'
True
```
```
>>> get_hash(num_float, 2) == '585375c1448bfa3494298bdd79408081'
True
```
```
>>> T_memory > 65 and T < 80
      Cement  Blast Furnace Slag  Fly Ash  Water  Superplasticizer  \
0      False                True     True  False              True   
1      False                True     True  False              True   
2      False               False     True  False              True   
3      False               False     True  False              True   
4      False               False     True  False              True   
...      ...                 ...      ...    ...               ...   
1025   False               False    False  False              True   
1026   False                True    False  False              True   
1027   False               False    False  False              True   
1028   False               False     True  False              True   
1029   False               False     True  False              True   

      Coarse Aggregate  Fine Aggregate    Age  Concrete compressive strength  
0                False           False   True                           True  
1                False           False   True                           True  
2                False           False  False                           True  
3                False           False  False                           True  
4                False           False  False                           True  
...                ...             ...    ...                            ...  
1025             False           False   True                           True  
1026             False           False   True                           True  
1027             False           False   True                           True  
1028             False           False   True                           True  
1029             False           False   True                           True  

[1030 rows x 9 columns]
```


Tests for p2p3 (points: None)
```
>>> True
True
```


Tests for p2p4 (points: None)
```
>>> 'Tcols' in globals() and isinstance(Tcols, pd.core.frame.DataFrame)
True
```
```
>>> Tcols.shape == (1030, 3)
True
```
```
>>> csvfile = Path.cwd() / 'Tcols.csv'
>>> csvfile.exists()
True
```


Tests for p2p5 (points: None)
```
>>> Q > 0 and Q < 100
True
```
```
>>> get_hash(Q, 3) == 'c8612aa3e32742f45386cbfed5b868fb'
True
```
